# CSCI 3155 : Homework 5

Evaluating expressions

__Name__:

Kyle Moe

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Artithmetic Expressions

$$\begin{array}{rcll}
\mathbf{AExpr} & := & Const(Int) \\
 & | & Ident(String) \\
 & | & Plus(\mathbf{AExpr}, \mathbf{AExpr}) \\
 & | & Minus(\mathbf{AExpr}, \mathbf{AExpr}) \\
 & | & Mult (\mathbf{AExpr}, \mathbf{AExpr}) \\
 & | & Neg (\mathbf{AExpr}) \\
\end{array}$$



$$ eval : AExpr \rightarrow Env \rightarrow Int $$

```scala
sealed trait AExpr
case class Const(n : Int) extends AExpr
case class Ident(s : String) extends AExpr
case class Plus( e1 : AExpr, e2 : AExpr) extends AExpr
case class Minus( e1 : AExpr, e2 : AExpr) extends AExpr
case class Mult( e1 : AExpr, e2 : AExpr) extends AExpr
case class Neg( e1 : AExpr) extends AExpr

def evalA( e : AExpr, env : Map[String,Int]) : Int = 
    e match {
        case Const(n)      => n 
        case Ident(s)      => env(s)
        case Plus(e1, e2)  => evalA(e1,env) + evalA(e2,env)
        case Minus(e1, e2) => evalA(e1,env) - evalA(e2,env)
        case Mult(e1, e2)  => evalA(e1,env) * evalA(e2,env)
        case Neg(e1)       => -evalA(e1,env)
    }

val myAExpr = Plus( Const(3), 
                   Mult( 
                       Ident("x"), 
                       Const(5)))
val myenv = Map("x" -> 4)
evalA( myAExpr, myenv)
```

Output:
```
defined [32mtrait[39m [36mAExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mNeg[39m
defined [32mfunction[39m [36mevalA[39m
[36mmyAExpr[39m: [32mPlus[39m = [33mPlus[39m([33mConst[39m([32m3[39m), [33mMult[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m5[39m)))
[36mmyenv[39m: [32mMap[39m[[32mString[39m, [32mInt[39m] = [33mMap[39m([32m"x"[39m -> [32m4[39m)
[36mres23_10[39m: [32mInt[39m = [32m23[39m
```

## Problem 1 : Boolean expressions

### Problem 1a (10 points)
Write a new type for boolean expressions, **BExpr**. This type includes everything necessary for simple boolean expressions: the boolean constants True and False, the boolean negation operation Not, the boolean binary operators And and Or, as well as ITE, the if then else expression. 

$$\begin{array}{rcll}
\mathbf{BExpr} & := & True \\
 & | & False \\
 & | & Not ( \mathbf{BExpr} )\\
 & | & And ( \mathbf{BExpr}, \mathbf{BExpr} ) \\
 & | & Or ( \mathbf{BExpr}, \mathbf{BExpr} ) \\
 & | & ITE(\mathbf{BExpr}, \mathbf{BExpr}, \mathbf{BExpr})\\
\end{array}$$

<!-- grade_id: cell-279792020634c1e2, points: 0, grade: false, solution: true -->
```scala
// Type def for Boolean Expression with True, False Base cases
sealed trait BExpr
case object True extends BExpr
case object False extends BExpr
case class Not(e : BExpr) extends BExpr
case class And(e1 : BExpr, e2 : BExpr) extends BExpr
case class Or(e1 : BExpr, e2 : BExpr) extends BExpr
case class ITE(e1 : BExpr, e2 : BExpr, e3 : BExpr) extends BExpr

```

Output:
```
defined [32mtrait[39m [36mBExpr[39m
defined [32mobject[39m [36mTrue[39m
defined [32mobject[39m [36mFalse[39m
defined [32mclass[39m [36mNot[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mITE[39m
```

<!-- grade_id: cell-4a4d1f097d5c1f68, points: 10, grade: true, solution: false -->
```scala
val bexpr1 : BExpr = True
val bexpr2 : BExpr = Not(bexpr1)
val bexpr3 : BExpr = And(Not(bexpr1), True)
val bexpr4 : BExpr = Or(Not(bexpr3), bexpr3)
val bexpr5 : BExpr = ITE(bexpr3, False, True)
val bexpr6 : BExpr = ITE(bexpr4, bexpr3, bexpr5)
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
[36mbexpr1[39m: [32mBExpr[39m = True
[36mbexpr2[39m: [32mBExpr[39m = [33mNot[39m(True)
[36mbexpr3[39m: [32mBExpr[39m = [33mAnd[39m([33mNot[39m(True), True)
[36mbexpr4[39m: [32mBExpr[39m = [33mOr[39m([33mNot[39m([33mAnd[39m([33mNot[39m(True), True)), [33mAnd[39m([33mNot[39m(True), True))
[36mbexpr5[39m: [32mBExpr[39m = [33mITE[39m([33mAnd[39m([33mNot[39m(True), True), False, True)
[36mbexpr6[39m: [32mBExpr[39m = [33mITE[39m(
  [33mOr[39m([33mNot[39m([33mAnd[39m([33mNot[39m(True), True)), [33mAnd[39m([33mNot[39m(True), True)),
  [33mAnd[39m([33mNot[39m(True), True),
  [33mITE[39m([33mAnd[39m([33mNot[39m(True), True), False, True)
)
```

### Problem 1b (20 points)
Write an evaluation function for boolean expressions using the above type.

$$\texttt{evalB} : BExpr\rightarrow Boolean$$

<!-- grade_id: cell-227d146a3f8c24a1, points: 0, grade: false, solution: true -->
```scala
// Evaluates a boolean expression starting with base case checks
def evalB( e : BExpr) : Boolean = {
    e match {
        case True => true 
        case False => false
        case Not(p) => !evalB(p) // Not of a boolean recursion
        case And(p1, p2) => evalB(p1) && evalB(p2) // And checks each 'lower branch' statements
        case Or(p1, p2) => evalB(p1) || evalB(p2) // Or like And above
        case ITE(p1, p2, p3) => if (evalB(p1)) evalB(p2) else evalB(p3) // ITE using inline if else and recursion
    }
}
```

Output:
```
defined [32mfunction[39m [36mevalB[39m
```

<!-- grade_id: cell-bd9565b6b5e47642, points: 20, grade: true, solution: false -->
```scala
assert(evalB(bexpr1) == true,  "test 1")
assert(evalB(bexpr2) == false, "test 2")
assert(evalB(bexpr3) == false, "test 3")
assert(evalB(bexpr4) == true,  "test 4")
assert(evalB(bexpr5) == true, " test 5")
assert(evalB(bexpr6) == false, "test 6")
passed(20)
```

Output:
```
*** Tests Passed (20 points) ***
```

## Problem 2: Combined expressions too.

### Problem 1a (10 points)

Write a new type **CExpr** for mixed arithmetic and boolean expressions. The new expressions should include everything from before, including the necessary components for both arithmetic boolean expressions, as well as the parts of useful combined expressions, Geq for numerical comparison, Eq for equality.


$$\begin{array}{rcll}
\mathbf{CExpr} & := & Const(Int) \\
 & | & True \\
 & | & False \\
 & | & Ident(String) \\
 & | & Plus(\mathbf{CExpr}, \mathbf{CExpr}) \\
 & | & Minus(\mathbf{CExpr}, \mathbf{CExpr}) \\
 & | & Mult (\mathbf{CExpr}, \mathbf{CExpr}) \\
 & | & Neg (\mathbf{CExpr}) \\
 & | & And ( \mathbf{CExpr}, \mathbf{CExpr} ) \\
 & | & Or ( \mathbf{CExpr}, \mathbf{CExpr} ) \\
 & | & Geq (\mathbf{CExpr}, \mathbf{CExpr}) \\
 & | & Eq (\mathbf{CExpr}, \mathbf{CExpr}) \\
 & | & ITE(\mathbf{CExpr}, \mathbf{CExpr}, \mathbf{CExpr})\\
\end{array}$$

NOTE: We will be reusing the constructor names in this notebook. Make sure to run the cell with the definitions you are using for each test!

<!-- grade_id: cell-c98d5002e417dc2a, points: 0, grade: false, solution: true -->
```scala
// Type def for General Expression in the order above
sealed trait CExpr
case class Const(n : Int) extends CExpr
case object True extends CExpr
case object False extends CExpr
case class Ident(s : String) extends CExpr
case class Plus( e1 : CExpr, e2 : CExpr) extends CExpr
case class Minus( e1 : CExpr, e2 : CExpr) extends CExpr
case class Mult( e1 : CExpr, e2 : CExpr) extends CExpr
case class Neg( e1 : CExpr) extends CExpr
case class And(e1 : CExpr, e2 : CExpr) extends CExpr
case class Not(e : CExpr) extends CExpr
case class Or(e1 : CExpr, e2 : CExpr) extends CExpr
case class Geq(e1 : CExpr, e2 : CExpr) extends CExpr
case class Eq(e1 : CExpr, e2 : CExpr) extends CExpr
case class ITE(e1 : CExpr, e2 : CExpr, e3 : CExpr) extends CExpr
```

Output:
```
defined [32mtrait[39m [36mCExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mobject[39m [36mTrue[39m
defined [32mobject[39m [36mFalse[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mNeg[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mNot[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mEq[39m
defined [32mclass[39m [36mITE[39m
```

<!-- grade_id: cell-105dea821bf91c90, points: 10, grade: true, solution: false -->
```scala
val expr1 : CExpr = Const(10)
val expr2 : CExpr = True
val expr3 : CExpr = False
val expr4 : CExpr = Ident("x")
val expr5 : CExpr = Plus(expr1, expr4)
val expr6 : CExpr = Minus(expr5, expr1)
val expr7 : CExpr = Mult(expr5, expr6)
val expr8 : CExpr = Neg(expr7)
val expr9 : CExpr = And(expr2,expr3)
val expr10 : CExpr = Or(expr2, expr9)
val expr11 : CExpr = Geq(expr1, expr4)
val expr12 : CExpr = Eq(expr1, expr4) 
val expr13 : CExpr = ITE(expr2, expr1, expr4)
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
[36mexpr1[39m: [32mCExpr[39m = [33mConst[39m([32m10[39m)
[36mexpr2[39m: [32mCExpr[39m = True
[36mexpr3[39m: [32mCExpr[39m = False
[36mexpr4[39m: [32mCExpr[39m = [33mIdent[39m([32m"x"[39m)
[36mexpr5[39m: [32mCExpr[39m = [33mPlus[39m([33mConst[39m([32m10[39m), [33mIdent[39m([32m"x"[39m))
[36mexpr6[39m: [32mCExpr[39m = [33mMinus[39m([33mPlus[39m([33mConst[39m([32m10[39m), [33mIdent[39m([32m"x"[39m)), [33mConst[39m([32m10[39m))
[36mexpr7[39m: [32mCExpr[39m = [33mMult[39m(
  [33mPlus[39m([33mConst[39m([32m10[39m), [33mIdent[39m([32m"x"[39m)),
  [33mMinus[39m([33mPlus[39m([33mConst[39m([32m10[39m), [33mIdent[39m([32m"x"[39m)), [33mConst[39m([32m10[39m))
)
[36mexpr8[39m: [32mCExpr[39m = [33mNeg[39m(
  [33mMult[39m(
    [33mPlus[39m([33mConst[39m([32m10[39m), [33mIdent[39m([32m"x"[39m)),
    [33mMinus[39m([33mPlus[39m([33mConst[39m([32m10[3
```

### Problem 2b (10 points)

Our eval function is no longer going to be able to return a single type since the expressions can return either Integers or Booleans. However we still want to be able to return a value for each expression. We might want to return either a number value or NumVal for integer or a boolean value or BinVal for booleans. If we run into an error in our expression we will want to return an Error.

$$\begin{array}{rcll}
\mathbf{Value} & := & NumVal(Int) \\
 & | & BinVal(Boolean) \\
 & | & Error \\
\end{array}$$

Write the **Value** type in scala:

<!-- grade_id: cell-f07a8638ca582024, points: 0, grade: false, solution: true -->
```scala
// Value type with Num and Bin constructors and Error object
sealed trait Value
case class NumVal(i : Int) extends Value
case class BinVal(b : Boolean) extends Value
case object Error extends Value
```

Output:
```
defined [32mtrait[39m [36mValue[39m
defined [32mclass[39m [36mNumVal[39m
defined [32mclass[39m [36mBinVal[39m
defined [32mobject[39m [36mError[39m
```

<!-- grade_id: cell-12035a00d3104015, points: 10, grade: true, solution: false -->
```scala
val a : Value = NumVal(10)
val b : Value = NumVal(-10)
val c : Value = BinVal(true)
val d : Value = Error
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
[36ma[39m: [32mValue[39m = [33mNumVal[39m([32m10[39m)
[36mb[39m: [32mValue[39m = [33mNumVal[39m([32m-10[39m)
[36mc[39m: [32mValue[39m = [33mBinVal[39m(true)
[36md[39m: [32mValue[39m = Error
```

### Problem 2c (10 points)

Environments store what values different variables represent. For arithmetic expressions the environment we used mapped variable names to integers. However in our new combined expressions integers aren't the only things that can be stored in the environment. What should the type of the environment be?

Scala allows us to define type aliases with the following syntax:

```type <Type Alias Name> = <Type Definition>```

If we wanted to introduce a new type in scala for bitstrings we could instead create a type alias for lists of booleans:

```type BitString = List[Boolean]```


What is the type of the environment for combined expressions?

<!-- grade_id: cell-c1a4b422ae999d40, points: 0, grade: false, solution: true -->
```scala
// We want our type Env to represent a map from a string (variable) to a Value obj
type Environment = Map[String, Value]
```

Output:
```
defined [32mtype[39m [36mEnvironment[39m
```

<!-- grade_id: cell-2126689af2f85656, points: 10, grade: true, solution: false -->
```scala
val emptyenv : Environment = Map()
val env1 : Environment = Map( "a" -> NumVal(10))
val env2 : Environment = Map( "b" -> BinVal(true))
val env3 : Environment = Map( "x" -> NumVal(3), "y" -> NumVal(4), "z" -> BinVal(true))
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
[36memptyenv[39m: [32mEnvironment[39m = [33mMap[39m()
[36menv1[39m: [32mEnvironment[39m = [33mMap[39m([32m"a"[39m -> [33mNumVal[39m([32m10[39m))
[36menv2[39m: [32mEnvironment[39m = [33mMap[39m([32m"b"[39m -> [33mBinVal[39m(true))
[36menv3[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mNumVal[39m([32m3[39m), [32m"y"[39m -> [33mNumVal[39m([32m4[39m), [32m"z"[39m -> [33mBinVal[39m(true))
```

### Problem 2d (20 points)

Write ``evalC``, the interpreter for evaluating combined arithmetic/boolean expressions:

$$ evalC : CExpr \rightarrow Environment \rightarrow Value $$

Suggestion: Start with one case at a time, and always consider the types of what being returned by every function.

<!-- grade_id: cell-d9c08124e6652827, points: 0, grade: false, solution: true -->
```scala
// Use pattern matching for evalC to distinguish the cases of the entered expression
def evalC( e : CExpr, env : Environment) : Value = {
    e match {
        case Const(i) => NumVal(i) // Handle constant int base case (return Value)
        case True => BinVal(true) // Base cases for BinVal
        case False => BinVal(false)
        case Ident(s) => if (env.contains(s)) env(s) else Error // Return env variable if exists otherwise Error
        
        // Pattern match the tuple of inputs to check that both are NatVals and add, othewise Error
        case Plus(p1, p2) => val e1 : Value = evalC(p1, env)
                             val e2 : Value = evalC(p2, env)
                             (e1, e2) match {
                                 case (NumVal(e1), NumVal(e2)) => NumVal(e1+e2)
                                 case _ => Error
                             }
        // Likewise with Minus and Mult, pattern match for Ints and return Value
        case Minus(p1, p2) => val e1 : Value = evalC(p1, env)
                              val e2 : Value = evalC(p2, env)
                              (e1, e2) match {
                                  case (NumVal(e1), NumVal(e2)) => NumVal(e1-e2)
                                  case _ => Error
                              }
        case Mult(p1, p2) => val e1 : Value = evalC(p1, env)
                             val e2 : Value = evalC(p2, env)
                             (e1, e2) match {
                                 case (NumVal(e1), NumVal(e2)) => NumVal(e1*e2)
                                 case _ => Error
                             }
        // Negate both Boolean and Int Values 
        case Neg(p) => val e : Value = evalC(p, env)
                       e match {
                           case NumVal(i) => NumVal(-i)
                           case BinVal(b) => BinVal(!b)
                           case _ => Error
                       }
        // With And, Or, and Geq run pattern matching to check boolean type and return BinVal of the conditional
        case And(b1, b2) => val e1 : Value = evalC(b1, env)
                            val e2 : Value = evalC(b2, env)
                            (e1, e2) match {
                                case (BinVal(p1), BinVal(p2)) => BinVal(p1 && p2)
                                case _ => Error
                            }
        case Or(b1, b2) => val e1 : Value = evalC(b1, env)
                           val e2 : Value = evalC(b2, env)
                           (e1, e2) match {
                               case (BinVal(e1), BinVal(e2)) => BinVal(e1 || e2)
                               case _ => Error
                           }
        case Geq(p1, p2) => val e1 : Value = evalC(p1, env)
                            val e2 : Value = evalC(p2, env)
                            (e1, e2) match {
                                case (NumVal(e1), NumVal(e2)) => BinVal(e1 >= e2)
                                case _ => Error
                            }
        // For Eq we want to return a Bool if inputs are Values but only true if 
        // inputs are of the same Value type and equivalent
        case Eq(p1, p2) => val e1 : Value = evalC(p1, env)
                           val e2 : Value = evalC(p2, env)
                           (e1, e2) match {
                               case (NumVal(e1), NumVal(e2)) => BinVal(e1 == e2)
                               case (BinVal(e1), BinVal(e2)) => BinVal(e1 == e2)
                               case (BinVal(e1), NumVal(e2)) => BinVal(false)
                               case (NumVal(e1), BinVal(e2)) => BinVal(false)
                               case _ => Error
                           }
        // ITE performs a simple boolean pattern match and runs designated input
        case ITE(p1, p2, p3) => val e1 : Value = evalC(p1, env)
                                e1 match {
                                    case BinVal(true) => evalC(p2, env)
                                    case BinVal(false) => evalC(p3, env)
                                    case _ => Error
                                }
    }
}
```

Output:
```
defined [32mfunction[39m [36mevalC[39m
```

<!-- grade_id: cell-fed8b7454159d4ed, points: 4, grade: true, solution: false -->
```scala
val e1 : CExpr = Const(10)
val e2 : CExpr = True
val emptyenv : Environment = Map()
assert(evalC( e1, emptyenv) == NumVal(10) , "test 1")
assert(evalC( e2, emptyenv) == BinVal(true), "test 2")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mCExpr[39m = [33mConst[39m([32m10[39m)
[36me2[39m: [32mCExpr[39m = True
[36memptyenv[39m: [32mEnvironment[39m = [33mMap[39m()
```

<!-- grade_id: cell-52acaa39646a2b6a, points: 4, grade: true, solution: false -->
```scala
val e1 : CExpr = Ident("x")
val e2 : CExpr = Ident("y")
val e3 : CExpr = Ident("z")
val env : Environment = Map("x"-> NumVal(10), "y" -> BinVal(true))
val emptyenv : Environment = Map()
assert(evalC( e1, env) == NumVal(10), "test 3")
assert(evalC( e2, env) == BinVal(true), "test 4")
assert(evalC( e1, emptyenv) == Error)
assert(evalC( e3, env) == Error)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mCExpr[39m = [33mIdent[39m([32m"x"[39m)
[36me2[39m: [32mCExpr[39m = [33mIdent[39m([32m"y"[39m)
[36me3[39m: [32mCExpr[39m = [33mIdent[39m([32m"z"[39m)
[36menv[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mNumVal[39m([32m10[39m), [32m"y"[39m -> [33mBinVal[39m(true))
[36memptyenv[39m: [32mEnvironment[39m = [33mMap[39m()
```

<!-- grade_id: cell-e010f23013305d3f, points: 4, grade: true, solution: false -->
```scala
val e1 : CExpr = Plus(Const(3), Mult(Ident("x"), Const(5)))
val e2 : CExpr = Minus(Mult(Ident("x"), Ident("x")), Mult(Ident("y"), Ident("y")))
val env1 : Environment = Map("x"-> NumVal(4), "y" -> NumVal(5))
val env2 : Environment = Map("x"-> BinVal(true))
assert(evalC( e1, env1) == NumVal(23), "test 5")
assert(evalC( e1, env2) == Error, "test 6")
assert(evalC( e2, env1) == NumVal(-9), "test 7")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mCExpr[39m = [33mPlus[39m([33mConst[39m([32m3[39m), [33mMult[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m5[39m)))
[36me2[39m: [32mCExpr[39m = [33mMinus[39m([33mMult[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"x"[39m)), [33mMult[39m([33mIdent[39m([32m"y"[39m), [33mIdent[39m([32m"y"[39m)))
[36menv1[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mNumVal[39m([32m4[39m), [32m"y"[39m -> [33mNumVal[39m([32m5[39m))
[36menv2[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mBinVal[39m(true))
```

<!-- grade_id: cell-f83148133d2d8471, points: 4, grade: true, solution: false -->
```scala
val e1 : CExpr = Neg(Const(10))
val e2 : CExpr = Neg(True)
val emptyenv : Environment = Map()
assert( evalC(e1, emptyenv) == NumVal(-10), "test 8")
assert( evalC(e2, emptyenv) == BinVal(false), "test 9")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mCExpr[39m = [33mNeg[39m([33mConst[39m([32m10[39m))
[36me2[39m: [32mCExpr[39m = [33mNeg[39m(True)
[36memptyenv[39m: [32mEnvironment[39m = [33mMap[39m()
```

<!-- grade_id: cell-dc66e65305e3b13e, points: 4, grade: true, solution: false -->
```scala
val e1 : CExpr = And(True, Or(Ident("x"), Neg(Ident("x"))))
val env : Environment = Map( "x" -> BinVal(false))
assert( evalC(e1, env) == BinVal(true), "test 10")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mCExpr[39m = [33mAnd[39m(True, [33mOr[39m([33mIdent[39m([32m"x"[39m), [33mNeg[39m([33mIdent[39m([32m"x"[39m))))
[36menv[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mBinVal[39m(false))
```

<!-- grade_id: cell-50821cc0a364b814, points: 4, grade: true, solution: false -->
```scala
val e1 : CExpr = Geq(Const(10), Const(5))
val e2 : CExpr = Geq(Const(5), Const(10))
val e3 : CExpr = Geq(Const(5), Const(5))
val e4 : CExpr = Geq(True, Const(5))
val emptyenv : Environment = Map()
assert( evalC(e1, emptyenv) == BinVal(true), "test 11")
assert( evalC(e2, emptyenv) == BinVal(false), "test 12")
assert( evalC(e3, emptyenv) == BinVal(true), "test 13")
assert( evalC(e4, emptyenv) == Error, "test 14")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mCExpr[39m = [33mGeq[39m([33mConst[39m([32m10[39m), [33mConst[39m([32m5[39m))
[36me2[39m: [32mCExpr[39m = [33mGeq[39m([33mConst[39m([32m5[39m), [33mConst[39m([32m10[39m))
[36me3[39m: [32mCExpr[39m = [33mGeq[39m([33mConst[39m([32m5[39m), [33mConst[39m([32m5[39m))
[36me4[39m: [32mCExpr[39m = [33mGeq[39m(True, [33mConst[39m([32m5[39m))
[36memptyenv[39m: [32mEnvironment[39m = [33mMap[39m()
```

<!-- grade_id: cell-523535459b6c64dc, points: 4, grade: true, solution: false -->
```scala
val e1 : CExpr = Eq( Const(10), Ident("x"))
val e2 : CExpr = Plus( Const(10), True)
val env1 : Environment = Map( "x" -> NumVal(10))
val env2 : Environment = Map( "x" -> NumVal(5))
val env3 : Environment = Map( "x" -> BinVal(true))
assert( evalC(e1, env1) == BinVal(true), "test 15")
assert( evalC(e1, env2) == BinVal(false), "test 16")
assert( evalC(e1, env3) == BinVal(false), "test 17")
assert( evalC(e2, env1) == Error, "test 18")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mCExpr[39m = [33mEq[39m([33mConst[39m([32m10[39m), [33mIdent[39m([32m"x"[39m))
[36me2[39m: [32mCExpr[39m = [33mPlus[39m([33mConst[39m([32m10[39m), True)
[36menv1[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mNumVal[39m([32m10[39m))
[36menv2[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mNumVal[39m([32m5[39m))
[36menv3[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mBinVal[39m(true))
```

<!-- grade_id: cell-0da9abb1ea5bbb4f, points: 4, grade: true, solution: false -->
```scala
val e1 : CExpr = ITE( Geq(Ident("x"), Ident("y")), Ident("x"), Ident("y"))
val env1 : Environment = Map( "x" -> NumVal(10), "y" -> NumVal(20))
val env2 : Environment = Map( "x" -> NumVal(5), "y" -> NumVal(4))
assert( evalC(e1, env1) == NumVal(20), "test 19")
assert( evalC(e1, env2) == NumVal(5), "test 20")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mCExpr[39m = [33mITE[39m([33mGeq[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)), [33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))
[36menv1[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mNumVal[39m([32m10[39m), [32m"y"[39m -> [33mNumVal[39m([32m20[39m))
[36menv2[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mNumVal[39m([32m5[39m), [32m"y"[39m -> [33mNumVal[39m([32m4[39m))
```

<!-- grade_id: cell-62fa522157d3d6af, points: 8, grade: true, solution: false -->
```scala
val e1 : CExpr =
    Mult( Ident("y"), 
        ITE( And( Geq(Ident("x"), Const(1)), Neg( Eq(Ident("x"), Const(0)))),
             Mult( 
                Minus( Ident("x"), Neg(Const(-1))), 
                Const(6)),
             Const(1))
        )
val env1 : Environment = Map("x" -> NumVal(5), "y" -> NumVal(5))
assert( evalC(e1, env1) == NumVal(120), "test 21")
passed(8)
```

Output:
```
*** Tests Passed (8 points) ***
[36me1[39m: [32mCExpr[39m = [33mMult[39m(
  [33mIdent[39m([32m"y"[39m),
  [33mITE[39m(
    [33mAnd[39m([33mGeq[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m)), [33mNeg[39m([33mEq[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m0[39m)))),
    [33mMult[39m([33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mNeg[39m([33mConst[39m([32m-1[39m))), [33mConst[39m([32m6[39m)),
    [33mConst[39m([32m1[39m)
  )
)
[36menv1[39m: [32mEnvironment[39m = [33mMap[39m([32m"x"[39m -> [33mNumVal[39m([32m5[39m), [32m"y"[39m -> [33mNumVal[39m([32m5[39m))
```