```scala

```

# CSCI 3155 Homework 2

Practice with Algebraic Data Types: Natural Numbers and Lists

__Name__:

Kyle Moe

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (20 Points) : Natural numbers

Natural numbers can be defined in scala using an abstract data type. First we create a type called ``Nat``, and then add two ways of constructing an element of type Nat. A Nat can either be a Zero, or the successor of another already defined Nat. With this type definition, we can construct elements of type Nat.

```scala
sealed trait Nat
case object Zero extends Nat
case class Succ(pred : Nat) extends Nat

val one   = Succ(Zero)
val two   = Succ(one)
val three = Succ(two)
val four  = Succ(three)
val five  = Succ(four)
```

Output:
```
defined [32mtrait[39m [36mNat[39m
defined [32mobject[39m [36mZero[39m
defined [32mclass[39m [36mSucc[39m
[36mone[39m: [32mSucc[39m = [33mSucc[39m(Zero)
[36mtwo[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m(Zero))
[36mthree[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))
[36mfour[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))
[36mfive[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))))
```

How are two natural numbers added together? The result is obtained by constructing a new natural number that is incremented from zero as many times as both of the operands combined.

```scala
def plus(x : Nat, y : Nat) : Nat = x match {
    case Zero    => y
    case Succ(x) => Succ(plus(x, y))
}

val ans = plus(two,three)
ans == five
```

Output:
```
defined [32mfunction[39m [36mplus[39m
[36mans[39m: [32mNat[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))))
[36mres44_2[39m: [32mBoolean[39m = true
```

### Problem 1a (10 points)
Define a function that multiplies two natural numbers:

<!-- grade_id: cell-07775b70f29b720c, points: 0, grade: false, solution: true -->
```scala
def mult(x : Nat, y : Nat) : Nat = {
    var result : Nat = Zero // Store the return var as a Nat
    if (x != Zero || y != Zero) { // Return Zero if either input is Zero
        // Define a recursive helper function to perform Nat multiplication
        def multRec(p : Nat, res: Nat) : Nat = p match { 
            case Zero => res // as p decrements, result += y
            case Succ(p) => multRec(p, plus(res,y))
        }
        // Call recsursion with with x as p
        result = multRec(x, result)
    }
    result
}
```

Output:
```
defined [32mfunction[39m [36mmult[39m
```

<!-- grade_id: cell-6cf23ea68d798237, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( mult(one, four) == four)
assert( mult(two, three) == Succ(five))
assert( mult(Zero, five) == Zero)
passed(10)
//END TEST
```

Output:
```
*** Tests Passed (10 points) ***
```

### Problem 1b (10 points)

Define a function called ``even`` that checks is a Nat is even, returning ``true`` if the number is even.

<!-- grade_id: cell-dcd74ba982b165c7, points: 0, grade: false, solution: true -->
```scala
// Apply pattern matching to develop a recusion with even/odd base cases
def even(n: Nat) : Boolean = n match {
    case Zero => true // even base case
    case Succ(Zero) => false // odd base case
    case Succ(Succ(prev)) => even(prev) // Subtract by two to maintain parity
}
```

Output:
```
defined [32mfunction[39m [36meven[39m
```

<!-- grade_id: cell-c18cb8cff48ffefb, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( even(Zero) == true)
assert( even(one) == false)
assert( even(five) == false)
assert( even(plus(five,five)) == true)
passed(10)
//END TEST
```

Output:
```
*** Tests Passed (10 points) ***
```

## Problem 2 (30 points) : Lists

### Problem 2a (10 points)
Write the definition of list of Nats in scala using using the cons construction we defined in class. If your definition is correct, the lists defined in the test should be created successfully.

<!-- grade_id: cell-05547163ae0a3a81, points: 0, grade: false, solution: true -->
```scala

sealed trait List // Define a List trait just as Nat
case object Empty extends List // Empty base case like Zero
case class Cons(n : Nat, lst: List) extends List // List constructor with two arguments

```

Output:
```
defined [32mtrait[39m [36mList[39m
defined [32mobject[39m [36mEmpty[39m
defined [32mclass[39m [36mCons[39m
```

<!-- grade_id: cell-c5a1cf21d594aaa1, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
val list1 = Empty
val list2 = Cons(five, Empty)
val list3 = Cons(three, Cons(one, Cons(four, Cons(one, Cons(five, Empty)))))
passed(10)
//END TEST
```

Output:
```
*** Tests Passed (10 points) ***
[36mlist1[39m: [32mEmpty[39m.type = Empty
[36mlist2[39m: [32mCons[39m = [33mCons[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))), Empty)
[36mlist3[39m: [32mCons[39m = [33mCons[39m(
  [33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))),
  [33mCons[39m(
    [33mSucc[39m(Zero),
    [33mCons[39m(
      [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))),
      [33mCons[39m([33mSucc[39m(Zero), [33mCons[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))), Empty))
    )
  )
)
```

### Problem 2b (5 points)
Write a ``List => Nat`` function called ``length`` that returns the length of a list as a Nat

<!-- grade_id: cell-198204e8fa90413b, points: 0, grade: false, solution: true -->
```scala
// Somewhat similar recursion to mult with helper function
def length(in : List) : Nat = {
    // helper function that increments size:Nat for each element
    def lengthRec(inLst : List, size : Nat) : Nat = {
        inLst match {
            case Empty => size // Empty list base case returns size
            case Cons(inNum, inLst) => lengthRec(inLst, Succ(size)) // recurse with incresed size and 'inside' list
        }
    }
    lengthRec(in, Zero) // Call helper with input and initial size of Zero
}
```

Output:
```
defined [32mfunction[39m [36mlength[39m
```

<!-- grade_id: cell-4851c80fe2d7a6b5, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( length(list1) == Zero )
assert( length(list2) == one )
assert( length(list3) == five )
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
```

### Problem 2c (5 points)
Write a ``sum`` function which computes the sum of a list and returns a Nat.

<!-- grade_id: cell-5f653fd80422308e, points: 0, grade: false, solution: true -->
```scala
// Similar to length this implements a helper function recursion
def sum(list : List) : Nat = {
    // Whereas length incremented by one for each element, sum adds the list 
    // number in Cons(inNum, inLst)
    def sumRec(inLst : List, amount : Nat) : Nat = {
        inLst match {
            case Empty => amount // Fully traversed list return sum
            case Cons(inNum, inLst) => sumRec(inLst, plus(inNum, amount)) // plus(inNum, amount) keeps track of sum
        }
    }
    sumRec(list, Zero) // Sum starts at Zero with input list
}
```

Output:
```
defined [32mfunction[39m [36msum[39m
```

<!-- grade_id: cell-dd1215a8749c149d, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( sum(list1) == Zero )
assert( sum(list2) == five )
assert( sum(list3) == Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero)))))))))))))) )
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
```

### Problem 2d (5 points)
Write a function called ``last`` that returns the last element of a list as a singleton list.

[1] => [1]

[1,2,3,4,5] => [5]

What should you do for empty lists?

<!-- grade_id: cell-a5e99da11725afdb, points: 0, grade: false, solution: true -->
```scala
// Another recursion with only one argument as we simply need a list element
def last(list : List) : List = {
    def lastRec(inner : List) : List = {
        // Check each inner list until reaching last element
        inner match {
            case Empty => inner // Handles empty list case
            case Cons(inNat, Empty) => inner // base case to return 
            case Cons(inNat, inList) => lastRec(inList) // Traverse list
        }
    }
    lastRec(list)
}
```

Output:
```
defined [32mfunction[39m [36mlast[39m
```

<!-- grade_id: cell-f3cd3bc0a702d90e, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( last(list1) == Empty )
assert( last(list2) == list2 )
assert( last(list3) == list2 )
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
```

### Problem 2e (5 points)
Write the ``countdown`` function: given a Nat x, return a list counting down from x to zero.

<!-- grade_id: cell-bcab1cbf421bcac1, points: 0, grade: false, solution: true -->
```scala
// Because we start with the list Cons(Zero,Empty) we can use our function, length,
// on the recursive list to get the Nat for each subsequent layer 
def countdown(n : Nat) : List = {
    def cdRec(left : Nat, list : List) : List = {
        left match {
            case Zero => list
            // Decrement our input Nat. Add addtional element with current list of value, length(list)
            case Succ(p) => cdRec(p, Cons(length(list), list))
        }
    }
    cdRec(n, Cons(Zero, Empty))
}
```

Output:
```
defined [32mfunction[39m [36mcountdown[39m
```

<!-- grade_id: cell-17ee23e31318b5e4, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( countdown(five) == 
    Cons( Succ(Succ(Succ(Succ(Succ(Zero))))),
    Cons( Succ(Succ(Succ(Succ(Zero)))),
    Cons( Succ(Succ(Succ(Zero))),
    Cons( Succ(Succ(Zero)), 
    Cons( Succ(Zero),
    Cons(Zero, 
         Empty))))))
)
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
```

```scala

```