```scala

```

# CSCI 3155 : Assignment 8

__Name__:

Kyle Moe

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 : Type Checking (100 points)

Write a function `typecheck` that typechecks a Lettuce program.

$$\texttt{typecheck} : Program \rightarrow Type$$ 

This function should take as input a lettuce _Program_, and return what _Type_ it is. If it doesn't typecheck it should return a special value _TypeError_. Use the Lettuce typing notes on moodle for details on how typechecking lettuce works.

The tests for this assigment are hidden, so you will need to write your own tests! Note that you are encouraged to collaborate and share your test ideas on Piazza. To get full points, your implementation needs to match the specification from the Lettuce typing notes on moodle.

```scala
sealed trait Type
case object NumType extends Type
case object BoolType extends Type
case class FunType( ty_a : Type, ty_r : Type) extends Type
case object TypeError extends Type

sealed trait Expr
case class  Const(n : Int) extends Expr
case class  Ident(s : String) extends Expr
case object True extends Expr
case object False extends Expr
case class  Plus( e1 : Expr, e2 : Expr) extends Expr
case class  Minus( e1 : Expr, e2 : Expr) extends Expr
case class  Mult( e1 : Expr, e2 : Expr) extends Expr
case class  Neg( e1 : Expr) extends Expr
case class  And( e1 : Expr, e2 : Expr) extends Expr
case class  Or( e1 : Expr, e2 : Expr) extends Expr
case class  Geq( e1 : Expr, e2 : Expr) extends Expr
case class  Eq( e1 : Expr, e2 : Expr) extends Expr
case class  ITE( ec : Expr, et : Expr, ef : Expr) extends Expr
case class  Let( name : String, ty_d : Type, ed : Expr, ei : Expr) extends Expr
case class  FunDef( arg : String, ty_a : Type, eb : Expr) extends Expr
case class  FunCall( ef : Expr, ea : Expr) extends Expr
case class  LetRec( f : String, ty_f : Type, x : String, ty_x : Type, eb : Expr, ei : Expr) extends Expr

sealed trait Program
case class Main(p : Expr) extends Program
```

Output:
```
defined [32mtrait[39m [36mType[39m
defined [32mobject[39m [36mNumType[39m
defined [32mobject[39m [36mBoolType[39m
defined [32mclass[39m [36mFunType[39m
defined [32mobject[39m [36mTypeError[39m
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mobject[39m [36mTrue[39m
defined [32mobject[39m [36mFalse[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mNeg[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mEq[39m
defined [32mclass[39m [36mITE[39m
defined [32mclass[39m [36mLet[39m
defined [32mclass[39m [36mFunDef[39m
defined [32mclass[39m [36mFunCall[39m
defined [32mclass[39m [36mLetRec[39m
defined [32mtrait[39m [36mProgram[39m
defined [32mcl
```

<!-- grade_id: cell-f8cd57266fd56cf2, points: 0, grade: false, solution: true -->
```scala
// By defining the following Context ADT we are able to store the type of each function and variable
sealed trait Context
case object Empty extends Context
case class Extend(key : String, t : Type, con : Context) extends Context

// Getter function for Context that determines if an Ident("X") corresponds to a known type
def get(c : Context, key : String) : Type = {
    c match {
        case Empty => TypeError 
        case Extend(k, t, con) => if (k == key) t else get(con, key)
    }
}

//Typecheck function for Main(Expr) kind of program
def typecheck(p : Program) : Type = {
    p match {
        //Ensure that our program is correctly constructed
        case Main(q) => {
            // We define a recursive helper function in order to make use of our defined Context
            def typeRec(i : Expr, con : Context) : Type = {
                i match {
                    case Const(n) => NumType 
                    case True => BoolType
                    case False => BoolType
                    case Ident(s) => get(con, s) // Depends on whether Identity has type in Context
                    
                    //Plus, Minus, and Mult are essentially the same and require that both args are
                    //NumType and if so returns NumType
                    case Plus(e1, e2) => {
                        if ((typeRec(e1, con) == NumType) && (typeRec(e2, con) == NumType)) NumType else TypeError
                    }
                    case Minus(e1, e2) => {
                        if ((typeRec(e1, con) == NumType) && (typeRec(e2, con) == NumType)) NumType else TypeError
                    }
                    case Mult(e1, e2) => {
                        if ((typeRec(e1, con) == NumType) && (typeRec(e2, con) == NumType)) NumType else TypeError
                    }
                    //Neg needs to be able to return return type if function or Expr type otherwise
                    case Neg(e) => typeRec(e, con) match {
                        case FunType(ta, tr) => tr
                        case BoolType => BoolType
                        case NumType => NumType
                        case TypeError => TypeError
                    }
                    // And and Or take BoolType args and returns BoolType
                    case And(e1, e2) => {
                        if ((typeRec(e1, con) == BoolType) && (typeRec(e2, con) == BoolType)) BoolType else TypeError
                    }
                    case Or(e1, e2) => {
                        if ((typeRec(e1, con) == BoolType) && (typeRec(e2, con) == BoolType)) BoolType else TypeError
                    }
                    //Geq and Eq take NumType and return BoolType
                    case Geq(e1, e2) => {
                        if ((typeRec(e1, con) == NumType) && (typeRec(e2, con) == NumType)) BoolType else TypeError
                    }
                    case Eq(e1, e2) => {
                        val t1 = typeRec(e1, con)
                        val t2 = typeRec(e2, con)
                        (t1, t2) match {
                            case (FunType(a, b), FunType(c, d)) => if (b == d) BoolType else TypeError
                            case (a, FunType(c, d)) => if (a == d) BoolType else TypeError
                            case (FunType(a, b), c) => if (b == c) BoolType else TypeError
                            case _ => if (t1 == t2) BoolType else TypeError
                        }
                        //if (typeRec(e1, con) == typeRec(e2, con)) BoolType else TypeError
                    }
                    // ITE requires that e1 is a BoolType conditional and that e2 and e3 are the same return type
                    case ITE(e1, e2, e3) => typeRec(e1, con) match {
                        case BoolType => {
                            val v2 = typeRec(e2, con)
                            val v3 = typeRec(e3, con)
                            if (v2 == v3) v2 else TypeError
                        }
                        case _ => TypeError
                    }
                    // We must evaluate the in expression so long as we know the let definition type in our Context
                    case Let(name, ty, ed, ei) => {
                        val newCon = Extend(name, ty, con)
                        typeRec(ei, newCon)
                    }
                    // FunDef returns FunType or Error but we must know argument type in Context to determine return type
                    case FunDef(a, tya, b) => {
                        val newCon = Extend(a, tya, con)
                        val tyb = typeRec(b, newCon)
                        tyb match {
                            case TypeError => TypeError
                            case _ => FunType(tya, tyb)
                        }
                    }
                    // When calling a function we want the function name to actually be associated with a
                    // FunType. From there ensure that argument type is appropriate as well
                    case FunCall(ef, ea) => {
                        val tyf = typeRec(ef, con)
                        val tya = typeRec(ea, con)
                        tyf match {
                            case FunType(ta, tr) => if (tya == ta) tr else TypeError
                            case _ => TypeError
                        }
                    }
                    // LetRec evaluates the function body in the context of the type of the function and its
                    // argument and from there ensures that the types of the function match the argument type
                    // and return type. The return type is defined by what returns from the in expression in
                    // when the function is included in the Context
                    case LetRec(f, tyf, a, tya, eb, ei) => {
                        val newCon1 = Extend(f, tyf, con)
                        val newCon2 = Extend(a, tya, newCon1)
                        val tyr = typeRec(eb, newCon2)
                        tyf match {
                            case FunType(ftya, ftyr) => {
                                if (ftya == tya && ftyr == tyr) {
                                    typeRec(ei, newCon1)
                                }else {
                                    TypeError
                                }
                                
                            }
                            case _ => TypeError
                        }
                        
                    }
                }
            }
            typeRec(q, Empty)
            
        }
        case _ => TypeError
    }
}
```

Output:
```
defined [32mtrait[39m [36mContext[39m
defined [32mobject[39m [36mEmpty[39m
defined [32mclass[39m [36mExtend[39m
defined [32mfunction[39m [36mget[39m
defined [32mfunction[39m [36mtypecheck[39m
```

<!-- grade_id: cell-d9b206afd7a24388, points: 20, grade: true, solution: false -->
```scala
val p1 : Program = Main( Const(10) )
assert( typecheck(p1) == NumType, "test 1")
passed(20)
println("Hidden Tests : up to 80 points remaining")
```

Output:
```
*** Tests Passed (20 points) ***
Hidden Tests : up to 80 points remaining
[36mp1[39m: [32mProgram[39m = [33mMain[39m([33mConst[39m([32m10[39m))
```

```scala
//Checks basic True/False
val p1 : Program = Main(True)
val p2 : Program = Main(False)
assert(typecheck(p1) == BoolType)
assert(typecheck(p2) == BoolType)
```

Output:
```
[36mp1[39m: [32mProgram[39m = [33mMain[39m(True)
[36mp2[39m: [32mProgram[39m = [33mMain[39m(False)
```

```scala
//Checks Plus
val p1 : Program = Main(Plus(Const(2), Const(-5)))
val p2 : Program = Main(Plus(Const(2), True))
val p3 : Program = Main(Plus(False, True))
assert(typecheck(p1) == NumType)
assert(typecheck(p2) == TypeError)
assert(typecheck(p3) == TypeError)
```

Output:
```
[36mp1[39m: [32mProgram[39m = [33mMain[39m([33mPlus[39m([33mConst[39m([32m2[39m), [33mConst[39m([32m-5[39m)))
[36mp2[39m: [32mProgram[39m = [33mMain[39m([33mPlus[39m([33mConst[39m([32m2[39m), True))
[36mp3[39m: [32mProgram[39m = [33mMain[39m([33mPlus[39m(False, True))
```

```scala
//Checks Minus
val p1 : Program = Main(Minus(Const(2), Const(-5)))
val p2 : Program = Main(Minus(Const(2), True))
val p3 : Program = Main(Minus(False, True))
assert(typecheck(p1) == NumType)
assert(typecheck(p2) == TypeError)
assert(typecheck(p3) == TypeError)
```

Output:
```
[36mp1[39m: [32mProgram[39m = [33mMain[39m([33mMinus[39m([33mConst[39m([32m2[39m), [33mConst[39m([32m-5[39m)))
[36mp2[39m: [32mProgram[39m = [33mMain[39m([33mMinus[39m([33mConst[39m([32m2[39m), True))
[36mp3[39m: [32mProgram[39m = [33mMain[39m([33mMinus[39m(False, True))
```

```scala
//Checks Mult
val p1 : Program = Main(Mult(Const(2), Const(-5)))
val p2 : Program = Main(Mult(Const(2), True))
val p3 : Program = Main(Mult(False, True))
assert(typecheck(p1) == NumType)
assert(typecheck(p2) == TypeError)
assert(typecheck(p3) == TypeError)
```

Output:
```
[36mp1[39m: [32mProgram[39m = [33mMain[39m([33mMult[39m([33mConst[39m([32m2[39m), [33mConst[39m([32m-5[39m)))
[36mp2[39m: [32mProgram[39m = [33mMain[39m([33mMult[39m([33mConst[39m([32m2[39m), True))
[36mp3[39m: [32mProgram[39m = [33mMain[39m([33mMult[39m(False, True))
```

```scala
//Checks Neg for NumType and BoolType
val p1 : Program = Main(Neg(Const(4)))
val p2 : Program = Main(Neg(False))
val p3 : Program = Main(Neg(Let( "f", FunType(NumType, NumType),
                     FunDef("x", NumType, Plus(Const(1), Ident("x"))),
                     FunCall( Ident("f"), Const(2)))))
assert(typecheck(p1) == NumType)
assert(typecheck(p2) == BoolType)
assert(typecheck(p3) == NumType)
```

Output:
```
[36mp1[39m: [32mProgram[39m = [33mMain[39m([33mNeg[39m([33mConst[39m([32m4[39m)))
[36mp2[39m: [32mProgram[39m = [33mMain[39m([33mNeg[39m(False))
[36mp3[39m: [32mProgram[39m = [33mMain[39m(
  [33mNeg[39m(
    [33mLet[39m(
      [32m"f"[39m,
      [33mFunType[39m(NumType, NumType),
      [33mFunDef[39m([32m"x"[39m, NumType, [33mPlus[39m([33mConst[39m([32m1[39m), [33mIdent[39m([32m"x"[39m))),
      [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m2[39m))
    )
  )
)
```

```scala
// Test ITE under multiple conditions as well as simple And and Or conditionals
val p1 : Program = Main(ITE(Geq(Const(1), Const(3)), True, False))
val p2 : Program = Main(ITE(Geq(Const(1), Const(3)), Const(0), False))
val p3 : Program = Main(ITE(Mult(Const(1), Const(4)), Const(0), Const(1)))
val p4 : Program = Main(And(Neg(False), True))
val p5 : Program = Main(Or(Neg(Const(-1)), True))
val p6 : Program = Main(ITE(False, Const(2), Mult(Const(3), Const(4))))

assert(typecheck(p1) == BoolType)
assert(typecheck(p2) == TypeError)
assert(typecheck(p3) == TypeError)
assert(typecheck(p4) == BoolType)
assert(typecheck(p5) == TypeError)
assert(typecheck(p6) == NumType)
```

Output:
```
[36mp1[39m: [32mProgram[39m = [33mMain[39m([33mITE[39m([33mGeq[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m3[39m)), True, False))
[36mp2[39m: [32mProgram[39m = [33mMain[39m([33mITE[39m([33mGeq[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m3[39m)), [33mConst[39m([32m0[39m), False))
[36mp3[39m: [32mProgram[39m = [33mMain[39m([33mITE[39m([33mMult[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m4[39m)), [33mConst[39m([32m0[39m), [33mConst[39m([32m1[39m)))
[36mp4[39m: [32mProgram[39m = [33mMain[39m([33mAnd[39m([33mNeg[39m(False), True))
[36mp5[39m: [32mProgram[39m = [33mMain[39m([33mOr[39m([33mNeg[39m([33mConst[39m([32m-1[39m)), True))
[36mp6[39m: [32mProgram[39m = [33mMain[39m([33mITE[39m(False, [33mConst[39m([32m2[39m), [33mMult[39m([33mConst[39m([32m3[39m), [33mConst[39m([32m4[39m))))
```

```scala
// Checks Let, FunDef, and FunCall
val p1 : Program = Main(Let( "f", FunType(NumType, NumType),
                     FunDef("x", NumType, Plus(Const(1), Ident("x"))),
                     FunCall( Ident("f"), Const(2))))
val p2 : Program = Main(Let("f", FunType(NumType, FunType(NumType, NumType)),
                   FunDef("x", NumType,
                     FunDef("y", NumType,
                       Plus(Ident("x"),Ident("y")))),
                   FunCall(FunCall(Ident("f"), Const(2)), Const(3))
                   ))
val p3 : Program = Main(Let("f", FunType(NumType, FunType(NumType, NumType)),
                   FunDef("x", NumType,
                     FunDef("y", NumType,
                       Plus(Ident("x"),Ident("y")))),
                   FunCall(FunCall(Ident("f"), False), Const(3))
                   ))
assert(typecheck(p1) == NumType)
assert(typecheck(p2) == NumType)
assert(typecheck(p3) == TypeError)
```

Output:
```
[36mp1[39m: [32mProgram[39m = [33mMain[39m(
  [33mLet[39m(
    [32m"f"[39m,
    [33mFunType[39m(NumType, NumType),
    [33mFunDef[39m([32m"x"[39m, NumType, [33mPlus[39m([33mConst[39m([32m1[39m), [33mIdent[39m([32m"x"[39m))),
    [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m2[39m))
  )
)
[36mp2[39m: [32mProgram[39m = [33mMain[39m(
  [33mLet[39m(
    [32m"f"[39m,
    [33mFunType[39m(NumType, [33mFunType[39m(NumType, NumType)),
    [33mFunDef[39m([32m"x"[39m, NumType, [33mFunDef[39m([32m"y"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))),
    [33mFunCall[39m([33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m2[39m)), [33mConst[39m([32m3[39m))
  )
)
[36mp3[39m: [32mProgram[39m = [33mMain[39m(
  [33mLet[39m(
    [32m"f"[39m,
    [33mFunType[39m(NumType, [33mFunType[39m(NumType, NumType)),
    [33mFunDef[39m([32m"x"[39m, NumT
```

```scala
// Additional FunDef checks for lambda functions
val p1 = Main( FunDef("g", NumType, Let("h", NumType, Const(1), Mult(Ident("g"), Ident("h")))) )
val p2 = Main( FunDef("f", NumType, Eq( Ident("f"), Ident("1"))) )
val p3 = Main( FunDef("e", NumType, Mult( Ident("e"), Ident("e"))) )
assert(typecheck(p1) == FunType(NumType, NumType))
assert(typecheck(p2) == TypeError)
assert(typecheck(p3) == FunType(NumType, NumType))
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m(
  [33mFunDef[39m(
    [32m"g"[39m,
    NumType,
    [33mLet[39m([32m"h"[39m, NumType, [33mConst[39m([32m1[39m), [33mMult[39m([33mIdent[39m([32m"g"[39m), [33mIdent[39m([32m"h"[39m)))
  )
)
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"f"[39m, NumType, [33mEq[39m([33mIdent[39m([32m"f"[39m), [33mIdent[39m([32m"1"[39m))))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"e"[39m, NumType, [33mMult[39m([33mIdent[39m([32m"e"[39m), [33mIdent[39m([32m"e"[39m))))
```

```scala
//Old Test -- Has FunDef within LetRec. The FunDef induces a typing error.
val p2 = Main(LetRec("fact", FunType(NumType, NumType), "x", NumType,
                        FunDef("x", NumType, 
                              ITE(Geq(Const(1), Ident("x")), 
                                 Const(1),
                                 Mult(Ident("x"), 
                                     FunCall(Ident("fact"),  
                                            Minus(Ident("x"), Const(1)))))),
                        FunCall(Ident("fact"), Const(5))))
// Young, sexy test without the redundant definition of "x" as NumType
val p1 = Main(LetRec("fact", FunType(NumType, NumType), "x", NumType,
                              ITE(Geq(Const(1), Ident("x")), 
                                 Const(1),
                                 Mult(Ident("x"), 
                                     FunCall(Ident("fact"),  
                                            Minus(Ident("x"), Const(1))))),
                        FunCall(Ident("fact"), Const(5))))
assert(typecheck(p1) == NumType)
typecheck(p2)
```

Output:
```
[36mp2[39m: [32mMain[39m = [33mMain[39m(
  [33mLetRec[39m(
    [32m"fact"[39m,
    [33mFunType[39m(NumType, NumType),
    [32m"x"[39m,
    NumType,
    [33mFunDef[39m(
      [32m"x"[39m,
      NumType,
      [33mITE[39m(
        [33mGeq[39m([33mConst[39m([32m1[39m), [33mIdent[39m([32m"x"[39m)),
        [33mConst[39m([32m1[39m),
        [33mMult[39m([33mIdent[39m([32m"x"[39m), [33mFunCall[39m([33mIdent[39m([32m"fact"[39m), [33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m))))
      )
    ),
    [33mFunCall[39m([33mIdent[39m([32m"fact"[39m), [33mConst[39m([32m5[39m))
  )
)
[36mp1[39m: [32mMain[39m = [33mMain[39m(
  [33mLetRec[39m(
    [32m"fact"[39m,
    [33mFunType[39m(NumType, NumType),
    [32m"x"[39m,
    NumType,
    [33mITE[39m(
      [33mGeq[39m([33mConst[39m([32m1[39m), [33mIdent[39m([32m"x"[39m)),
      [33mConst[39m([32m1[39m),
      [33mMult[39m([33mIdent[39m([
```

<!-- grade_id: cell-31198cd330553090, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-69b68741c9fda826, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-d15f5cb6ed1607da, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-d8ea3308b64e150d, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-83408b1f20330fbc, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-8f34165e14575453, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-bb4cebe2260d7612, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-c4f5dcf0bf8fdca3, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-8037f9727ae33dd4, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-b59243a5817202d2, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-ede2ed89a8f56746, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-90690da2df5d6da5, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-288a52329f6c5428, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-7d0654077dff6d8b, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-d0489819cc1d8319, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-533d45c0493a5630, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-97f81d04838e2546, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-f91107fcf26e2240, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-3ba6c8052c148fff, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-36a0b90b6303d35e, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-89108e7186bf2e7e, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-e0c1267b4358e17e, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-7755f725e7af3e5c, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-a16534250c0ac8f9, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-6089a04546621af9, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-ea1c6a51893fd4d1, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-b3233ea6f1403203, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-0934c6879471b3a3, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-750499f6519c9cb8, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-b9646dcc63b28149, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-3f118b29e49adbec, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-150fee1089cc254b, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-f7fad77d24bd342e, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

```scala

```